---
title: "Sitzung 3: Konfirmatorische Faktorenanalyse"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
    includes:
      after_body: footer.html
      in_header: header.html
runtime: shiny_prerendered
---

```{r setup, include = FALSE}
library(learnr)
library(gradethis)
library(shiny)
library(fontawesome)
library(lavaan)

data(conspiracy, package = 'PsyMSc1')
knitr::opts_chunk$set(exercise.checker = gradethis::grade_learnr)
```

```{r, echo = FALSE}
abbrev <- function(X, begin = 'Latent Variables', end = NULL, ellipses = 'both', shift = 2, ...) {

  tmp <- capture.output(lavaan::summary(X,...))

  if (is.null(begin)) begin <- 1
  else begin <- grep(begin, tmp, fixed = TRUE)[1]
  if (is.null(end)) end <- length(tmp)-shift
  else end <- grep(end, tmp, fixed = TRUE)[grep(end, tmp, fixed = TRUE) > begin][1]-shift

  if (ellipses == 'both') {
    cat('[...]\n', paste(tmp[begin:end], collapse = '\n'), '\n[...]\n')
  }
  if (ellipses == 'top') {
    cat('[...]\n', paste(tmp[begin:end], collapse = '\n'))
  }
  if (ellipses == 'bottom') {
    cat(paste(tmp[begin:end], collapse = '\n'), '\n[...]\n')
  }
  if (ellipses == 'none') {
    cat(paste(tmp[begin:end], collapse = '\n'))
  }
}
```


## Einleitung

In der letzten Sitzung wurden faktoranalytische Verfahren für Datenexploration behandelt. Die Ergebnisse der EFA sind datengesteuert: welche Items welchen Faktoren zugeordnet werden, wie viele Faktoren genutzt werden, wie stark der Zusammenhang zwischen Item und Faktor ist, das alles sind Dinge, die aus den Daten heraus entschieden werden. In dieser Sitzung betrachten wir das Vorgehen, wenn in der Faktorenanalysen von einem konkreten, theoretisch fundierten Modell ausgegangen wird und dieses anhand empirischer Daten geprüft werden soll. Ganz im Popper'schen Sinn lässt sich nur durch ein solches Vorgehen wissenschaftlich Erkenntnis gewinnen.

Für die Umsetzung der CFA nutzen wir erneut `lavaan` - gucken Sie gerne noch einmal in die Unterlagen aus der 1. Sitzung um sich zu vergegenwärtigen, wie hier die Modellierung erfolgt. Grob gesagt gehen wir in `lavaan` in drei Schritten vor:

  1. Modellsyntax schreiben und in einem Objekt ablegen
  2. Modell schätzen und Ergebnisse in Objekt ablegen
  3. Ergebnisse im Objekt aus Schritt 2 inspizieren

In der ersten Sitzung hatten wir gesehen, dass in `lavaan` anhand weniger Begriffe das Pfadmodell in Syntax übersetzt wird, um unser Modell schätzbar zu machen. Diese Begriffe waren

| Bezeichnung | Befehl | Bedeutung |
|:-----------:|:------:|:---------:|
| Regression | `~` | wird vorhergesagt durch |
| Kovarianz | `~~` | kovariiert mit |
| Intercept | `~1` | wird auf 1 regressiert |
| Faktorladung | `=~` | wird gemessen durch |
| Formative Faktoren | `<~` | wird konstruiert durch |
| Schwellenparameter | `|t...` | Schwelle Nummer ... |

Bisher haben wir nur mit `~`, `~1` und `~~` für Regressionen, Intercepts und Residualvarianzen gearbeitet. Heute kommen Faktorladungen dazu.

Wie schon für letzten beiden Sitzungen, können Sie das komplette R-Skript [`r fontawesome::fa("download")` hier herunterladen](https://raw.githubusercontent.com/martscht/PsyMSc1/master/inst/scripts/cfa.R).

## Datensatz

Konkret wird es in dieser Sitzung um die Faktorstruktur verschwörungstheoretischer Überzeugungen gehen. Dafür nutzen wir Daten aus der Erhebung zur Validierung der *Generic Conspiracist Beliefs Scale* (GCBS; Brotherton, French, & Pickering, 2013). Die Daten finden Sie öffentlich zugänglich auf der [Open Psychometrics Website](http://openpsychometrics.org/_rawdata/GCBS.zip). Natürlich können Sie dort den [Fragebogen auch selbst ausfüllen](https://openpsychometrics.org/tests/GCBS/), wenn Sie möchten. Der Fragebogen besteht aus insgesamt 15 Aussagen, die Probanden jeweils von 1 ("definitely not true") bis 5 ("definitely true") hinsichtlich ihres Wahrheitsgehalts einschätzen sollen.

| Nr. | Facette | Itemformulierung |
|:-:| --- | ---------------------- |
| 1 | GM | The government is involved in the murder of innocent citizens and/or well-known public figures, and keeps this a secret |
| 2 | GC | The power held by heads of state is second to that of small unknown groups who really control world politics |
| 3 | EC | Secret organizations communicate with extraterrestrials, but keep this fact from the public |
| 4 | PW | The spread of certain viruses and/or diseases is the result of the deliberate, concealed efforts of some organization |
| 5 | CI | Groups of scientists manipulate, fabricate, or suppress evidence in order to deceive the public |
| 6 | GM | The government permits or perpetrates acts of terrorism on its own soil, disguising its involvement |
| 7 | GC | A small, secret group of people is responsible for making all major world decisions, such as going to war |
| 8 | EC | Evidence of alien contact is being concealed from the public |
| 9 | PW | Technology with mind-control capacities is used on people without their knowledge |
| 10 | CI | New and advanced technology which would harm current industry is being suppressed |
| 11 | GM | The government uses people as patsies to hide its involvement in criminal activity |
| 12 | GC | Certain significant events have been the result of the activity of a small group who secretly manipulate world events |
| 13 | EC | Some UFO sightings and rumors are planned or staged in order to distract the public from real alien contact |
| 14 | PW | Experiments involving new drugs or technologies are routinely carried out on the public without their knowledge or consent |
| 15 | CI | A lot of important information is deliberately concealed from the public out of self-interest |

Die 15 Aussagen können auf theoretischer Grundlage fünf verschiedenen Facetten zugeordnet werden, die in der Tabelle mit ihren jeweiligen Abkürzungen eingetragen sind:

  - GM: Government malfeasance
  - GC: Malevolent global conspiracies
  - EC: Extraterrestrial cover-up
  - PW: Personal Well-Being
  - CI: Control of Information

In dieser Sitzung werden wir untersuchen inwiefern diese theoretische Zuordnung einer empirischen Prüfung standhält.

Zunächst können wir den Datensatz laden und ein bisschen genauer betrachten. Wir haben den Datensatz in diesem Paket integriert, sodass Sie ihn über die `data`-Funktion aufrufen können:

```{r, eval = FALSE}
data(conspiracy, package = 'PsyMSc1')
```

Wir können den Datensatz einfach direkt aufrufen um uns einen Überblick zu verschaffen:

```{r}
conspiracy
```

Der Datensatz ist allerdings etwas groß um so einen wirklich guten Überblick zu erhalten. Im Folgenden `R`-Abschnitt können Sie die üblichen Funktionen nutzen, um ein paar generelle Eigenschaften des Datensatzes in Erfahrung zu bringen - z.B. generelle Struktur, Namen der Variablen, Anzahl der Personen. Falls Sie ein paar Vorschläge bezüglich der Funktionen benötigen, nutzen Sie einfach den *`r fontawesome::fa('lightbulb')` Hints* Button!

```{r data_overview, exercise = TRUE, exercise.eval = FALSE}

```

```{r data_overview-hint-1}
# Allgemeine Struktur des Datensatzes
str(conspiracy)
```

```{r data_overview-hint-2}
# Namen aller Variable im Datensatz
names(conspiracy)
```

```{r data_overview-hint-3}
# Anzahl der Zeilen (Personen) und Spalten (Variablen)
nrow(conspiracy)
ncol(conspiracy)
```

```{r data_overview-hint-4}
# Die ersten 6 Zeilen ansehen
head(conspiracy)
```

## Einfaktorielles Modell

In der Konzeption der verschwörungstheoretischen Überzeugungen nach Brotherton, French und Pickering (2013) wird behauptet, dass die Items 2, 7 und 12 ein gemeinsames Konstrukt - nämlich "Malevolent global conspiracies" - erheben. Hier noch einmal die genaue Formulierung der drei Aussagen:

| Nr. | Itemformulierung |
|:-:| ---------------------- |
| 2 | The power held by heads of state is second to that of small unknown groups who really control world politics |
| 7 | A small, secret group of people is responsible for making all major world decisions, such as going to war |
| 12 | Certain significant events have been the result of the activity of a small group who secretly manipulate world events |

Es wird also behauptet, dass es eine grundlegende Eigenschaft von Personen gibt, die das Ausmaß bestimmt in dem Personen diese Aussagen als wahr einschätzen. Mithilfe der CFA können wir versuchen, diese Behauptung zu prüfen.

Als Pfadmodell dargestellt, sieht das zunächst so aus:

<img src="https://raw.githubusercontent.com/martscht/PsyMSc1/Sitzung3/inst/tutorials/CFA/images/cfa1.png" width="33%"/>

### Modellsyntax

Wie wir es in der ersten Sitzung für die Regression gesehen haben, können wir dieses Pfaddiagramm in `lavaan` Modellsyntax übersetzen. Zunächst haben wir eine latente Variable, die wir nach Belieben benennen können. Weil es sich sinnvoll anbietet, nenne ich sie hier `GC` für "global conspiracy". Diese latente Variable lädt auf drei manifeste Variablen. Bei diesen Variablen habe ich nicht die gleichen Freiheiten in der Benennung, weil es sich hier um manifeste Variablen handelt, die im Datensatz bereits existieren. Daher muss ich hier die auch die Namen verwenden, die im Datensatz für die Variablen genutzt werden. Mit dem `=~` wird die Beziehung zwischen einer latenten Variable (links) und beliebig vielen manifesten Variablen (rechts) dargestellt. Diese Beziehung wird allgemein als Faktorladung bezeichnet. Mehrere manifeste Variablen, die der gleichen latenten Variable zugeordnet werden, können mit einem `+` verbunden werden. Für die Faktorladungen sieht der Modellsyntax-Abschnitt also so aus:

```{r, eval = FALSE}
'GC =~ Q2 + Q7 + Q12'
```

Wie schon in der Sitzung zur Regression müssen wir außerdem explizit Residualvarianzen anfordern, weil das Modell sonst behaupten würde, dass die latente Variable eine Varianzaufklärung von 100% an unseren manifesten Variablen leistet. Obwohl das zwar außerordentlich schön für die Entwickler des Fragebogens wäre, wenn dem so wäre, ist beinahe jede manifeste Variable in der psychologischen Forschung mit einem gewissen Maß an Messfehler behaftet. Dieser wird in der CFA als Residualvarianz abgebildet. Daher erweitert sich unser Modell um die folgenden drei Zeilen:

```{r, eval = FALSE}
'# Faktorladungen
GC =~ Q2 + Q7 + Q12

# Residualvarianzen
Q2 ~~ Q2
Q7 ~~ Q7
Q12 ~~ Q12'
```

Um das Modell ein wenig zu gliedern, habe ich außerdem mit dem `#` Kommentare in die Modellsyntax eingefügt, die die einzelnen Abschnitte voneinander abgrenzen. Auch dieses Modell ist noch nicht ganz vollständig. Weil die Varianz der latenten Variable `GC` nicht explizit angesprochen wird, wird - wie schon bei den Residuen - per Voreinstellung davon ausgegangen, dass diese Varianz 0 ist. Um Varianz auf der latenten Variable zuzulassen, müssen wir also auch diese explizit in das Modell aufnehmen:

```{r, eval = FALSE}
'# Faktorladungen
GC =~ Q2 + Q7 + Q12

# Residualvarianzen
Q2 ~~ Q2
Q7 ~~ Q7
Q12 ~~ Q12

# Latente Varianz
GC ~~ GC'
```

Dieses Modell müssen wir erneut einem Objekt zuweisen, damit wir es gleich in der `lavaan`-Funktion ansprechen können. In einer glorreichen Demonstration meines Einfallsreichtums, habe ich dieses Objekt `mod1` benannt - Modell 1.

```{r model1}
mod1 <- '# Faktorladungen
GC =~ Q2 + Q7 + Q12

# Residualvarianzen
Q2 ~~ Q2
Q7 ~~ Q7
Q12 ~~ Q12

# Latente Varianz
GC ~~ GC'
```

### Modellschätzung und -ergebnisse

Sobald wir das Modell erstellt haben folgen alle weiteren Schritte dem gleichen Prinzip, das wir in der ersten Sitzung für Regressionen gesehen haben. Zunächst wenden wir das Modell mit dem `lavaan`-Befehl auf den Datensatz an. Dieser Datensatz muss dabei alle Variablen enthalten, die wir als manifeste Variablen in unserem Modell angesprochen haben.

```{r fit1}
fit1 <- lavaan(mod1, conspiracy)
```

Die entstehende Warnung stellt ein ernstzunehmendes Problem dar und wird Ihnen wahrscheinlich häufiger begegnen, wenn Sie Dinge mit `lavaan` ausprobieren. Die Warnung gibt Ihnen zunächst darüber Auskunft wo das Problem aufgetreten ist - in diesem Fall in einer `lavaan`-internen Funktion namens `lav_model_vcov` - was aber nicht wirklich von Relevanz ist. Viel wichtiger sind die drei Informationen, die im Text der Warnmeldung enthalten sind:

  1. `Could not compute standard errors!` Es konnten keine Standardfehler bestimmt werden, was bedeutet, dass Sie keinerlei Inferenzstatistik betreiben können.
  2. `The information matrix could not be inverted.` Der technische Grund, aus dem keine Standardfehler bestimmt werden konnten.
  3. `This may be a symptom that the model is not identified.` Die wahrscheinliche Ursache, an der Sie ansetzen sollten um dieses Problem zu beheben.

Wir können uns die entstandenen Ergebnisse angucken, um herauszufinden worin genau das Problem liegen könnte:

```{r}
summary(fit1)
```

Die folgende Zeile des Outputs zeigt uns, dass `lavaan` mit seiner Einschätzung vollkommen Recht hatte:

```{r, echo = FALSE}
abbrev(fit1, 'Degrees of freedom', 'P-value', shift = 1)
```

In diesem Fall sind die Freiheitsgrade negativ - das Modell ist also *global unteridentifiziert*. Bevor Sie jetzt auf den *Continue* Knopf drücken um herauszufinden, wie man dieses Problem behebt, überlegen Sie kurz anhand Ihres Wissens aus dem Theorie-Teil was genau das Problem sein könnte und wie Sie es beheben würden, wenn die Sitzung hier einfach vorbei wäre.

### Modellidentifikation

Sie wissen inzwischen bestimmt, wo sich das Problem in diesem Fall verbirgt. Aber haben Sie trotzdem etwas Geduld - wir werden uns hier ein generelles Vorgehen in `lavaan` ansehen, mit dem man solche Probleme auch in komplizierteren Situationen lösen kann.

Zunächst wissen Sie aus der Theorie-Sitzung, dass die Anzahl bekannter Informationen sich mit folgender Gleichung berechnen lässt:

$n_{info} = \frac{p(p + 1)}{2}$

wobei $p$ die Anzahl von manifesten Variablen im Modell ist. In unserem Fall also

```{r}
3 * (3 + 1) / 2
```

Wir können uns auch genau angucken, welche `r 3 * (3 + 1) / 2` Informationen wir haben, indem wir uns die Kovarianzmatrix aus unseren Daten (*sample statistics*) ansehen:

```{r}
inspect(fit1, 'sampstat')
```

Hier können wir einfach durchzählen:

```{r, echo = FALSE}
tmp <- inspect(fit1, 'sampstat')
tmp$cov[lower.tri(tmp$cov, diag = TRUE)] <- 1:6
tmp
```

Die zu schätzenden Parameter ($n_{par}$) können wir im Pfaddiagramm abzählen:

<img src="https://raw.githubusercontent.com/martscht/PsyMSc1/Sitzung3/inst/tutorials/CFA/images/cfa2.png" width="33%"/>

Wenn wir uns nicht damit abmühen möchten die Parameter selbst zu zählen, können wir die Anzahl natürlich auch einfach von `lavaan` bekommen:

```{r}
inspect(fit1, 'npar')
```

Welche Parameter das genau sind, die hier geschätzt werden, können wir auch mit `inspect` erfahren:

```{r}
inspect(fit1, 'free')
```

Hier sehen wir, wie in `lavaan` die Modell Matrizen dargestellt werden:

  - `$lambda`: die Matrix der Faktorladungen ($\Lambda$)
  - `$theta`: die Residualkovarianzmatrix ($\Theta$)
  - `$psi`: die latente Kovarianzmatrix ($\Psi$)

Das Ergebnis von `inspect` ist hier eine Liste dieser drei Matrizen. Wenn wir diese Liste in einem Objekt ablegen, können wir die Matrizen mit dem `$` auch einzeln inspizieren:

```{r}
free <- inspect(fit1, 'free')
free$theta
```

In `theta` ist also die Residualvarianz von Q7 der insgesamt 5. Parameter des Modells.

Im Theorie-Abschnitt haben Sie gehört, dass in Modellen mit latenten Variablen immer entweder eine Faktorladung oder die latente Varianz fixiert werden muss, um der latenten Variable eine Skala zu geben. Entweder wird eine Faktorladung auf 1 gesetzt um so die Skala der manifesten Variable auf die latente Variable anzuwenden, oder die latente Varianz wird auf 1 fixiert, um die latente Variable in $z$-Werten interpretierbar zu machen. In der EFA ist diese zweite Variante extrem verbreitet. In der CFA wird Ihnen die erste Variante häufiger begegnen.

Um die erste Faktorladung auf 1 zu fixieren, können wir das `lavaan`-Modell wie folgt ergänzen:

```{r}
mod1 <- '# Faktorladungen
GC =~ 1*Q2 + Q7 + Q12

# Residualvarianzen
Q2 ~~ Q2
Q7 ~~ Q7
Q12 ~~ Q12

# Latente Varianz
GC ~~ GC'
```

Mit `1*` legen wir den konkreten Wert der Faktorladung fest. Das gleiche Prinzip gilt in `lavaan` für alle Modellparameter. Probieren Sie unten aus, das Modell so zu schreiben, dass die Faktorladungen frei sind, aber die latente Varianz auf 1 fixiert ist!

```{r std_lv, exercise = TRUE, exercise.eval = FALSE}
mod1b <- '# Faktorladungen
GC =~ Q2 + Q7 + Q12

# Residualvarianzen
Q2 ~~ Q2
Q7 ~~ Q7
Q12 ~~ Q12

# Latente Varianz
GC ~~ GC'

mod1b
```

```{r std_lv-check}
tmp <- '# Faktorladungen
GC =~ Q2 + Q7 + Q12

# Residualvarianzen
Q2 ~~ Q2
Q7 ~~ Q7
Q12 ~~ Q12

# Latente Varianz
GC ~~ 1*GC'

grade_result(
  fail_if(~ (!grepl('~~1*Q', gsub(' ', '', .result))), 'Die Residualvarianz brauchen Sie nicht zu fixieren. Diese können ruhig geschätzt werden.'),
  fail_if(~ (!grepl('=~1*Q', gsub(' ', '', .result))), 'Wenn Sie die latente Varianz fixieren brauchen Sie nicht gleichzeitig die Faktorladung zu fixieren. Eins von beiden ist ausreichend.'),
  pass_if(~ identical(gsub(' ', '', .result), gsub(' ', '', tmp))),
  correct = 'Sehr gut! In diesem Fall wird die latente Variable z-standardisiert.',
  incorrect = 'Leider falsch.',
  glue_correct = '{.correct}',
  glue_incorrect = '{.incorrect} {.message}')
```

### Modellschätzung und Ergebnisse - Zweiter Versuch



## Literatur
[Eid, M., Gollwitzer, M., & Schmitt, M. (2017).](https://hds.hebis.de/ubffm/Record/HEB366849158) *Statistik und Forschungsmethoden* (5. Auflage, 1. Auflage: 2010). Weinheim: Beltz.

[Brotherton, R., French, C. C., & Pickering, A. D. (2013)](https://www.frontiersin.org/articles/10.3389/fpsyg.2013.00279/full). Measuring belief in conspiracy theories: the generic conspiracist beliefs scale. *Frontiers in Psychology*, 4, 279, DOI: 10.3389/fpsyg.2013.00279.

<small> *Blau hinterlegte Autorenangaben führen Sie direkt zur Ressource.*
